# CS 6290: Privacy-enhancing Technologies
## Assignment 2

(Acknowledgement: this assignment is adapted from [Homework 7 for UVM CS 3110, developed by Joseph P. Near](https://github.com/jnear/cs3110-data-privacy/blob/main/homework/CS3110_HW_7.ipynb).)

```python
# Install dependencies (Suppose you are using Python3)
import sys
import platform

# Detect OS and adjust pip install command
if platform.system() in ["Linux", "Darwin"]:  # Darwin is macOS
    pip_cmd = f"{sys.executable} -m pip install --break-system-packages"
else:  # Windows or others
    pip_cmd = f"{sys.executable} -m pip install"

# Install dependencies
!{pip_cmd} pandas
!{pip_cmd} numpy
!{pip_cmd} scipy
!{pip_cmd} matplotlib
```

Output:
```
Requirement already satisfied: pandas in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (2.2.3)
Requirement already satisfied: numpy>=1.26.0 in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (from pandas) (1.26.4)
Requirement already satisfied: python-dateutil>=2.8.2 in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (from pandas) (2.9.0)
Requirement already satisfied: pytz>=2020.1 in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (from pandas) (2024.1)
Requirement already satisfied: tzdata>=2022.7 in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (from pandas) (2024.2)
Requirement already satisfied: six>=1.5 in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (from python-dateutil>=2.8.2->pandas) (1.16.0)
Requirement already satisfied: numpy in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (1.26.4)
Requirement already satisfied: scipy in /opt/anaconda3/envs/cityu/lib/python3.12/site-packages (1.14.1)
Requirement already satisfied: 
```

```python
# Load the data and libraries
import pandas as pd
import numpy as np
import random
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def laplace_mech_vec(vec, sensitivity, epsilon):
    return [v + np.random.laplace(loc=0, scale=sensitivity / epsilon) for v in vec]

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def gaussian_mech_vec(vec, sensitivity, epsilon, delta):
    return [v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)
            for v in vec]

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('./adult_with_pii.csv')
```


## Prerequisites

- Before diving into the programming, we suggest reviewing the following resources to refresh or build your understanding:
  - [Differential Privacy](https://programming-dp.com/ch3.html) – Revisit the basics of differential privacy and the Laplace Mechanism.
  - [Properties of Differential Privacy](https://programming-dp.com/ch4.html) – Learn about Sequential Composition in DP.
  - [The Gaussian Mechanism](https://programming-dp.com/ch6.html#the-gaussian-mechanism) – Explore the Gaussian Mechanism.
  - [Variants of Differential Privacy](https://programming-dp.com/ch8.html) – Understand Rényi DP.

- For those eager to deepen their understanding of the theoretical foundations of differential privacy, consider reading:
  - [Chapters 2 & 3 of *The Algorithmic Foundations of Differential Privacy*](https://www.cis.upenn.edu/~aaroth/Papers/privacybook.pdf) by Cynthia Dwork and Aaron Roth.
  
- We assume you have basic knowledge about `numpy` and `pandas`. For those unfamiliar with these libraries, please visit the crash course provided by Kaggle https://www.kaggle.com/learn/pandas.

## Range Queries

A *range query* counts the number of rows in the dataset which have a value lying in a given range. For example, "how many participants are between the ages of 21 and 33?" is a range query. A *workload* of range queries is just a list of range queries. The code below generates 100 random range queries over ages in the adult dataset.

```python
def range_query(df, col, a, b):
    return len(df[(df[col] >= a) & (df[col] < b)])

random_lower_bounds = [random.randint(1, 70) for _ in range(100)]
random_workload = [(lb, random.randint(lb, 100)) for lb in random_lower_bounds]
real_answers = [range_query(adult, 'Age', lb, ub) for (lb, ub) in random_workload]
print('First 5 queries: ', random_workload[:5])
```

Output:
```
First 5 queries:  [(18, 95), (62, 79), (23, 25), (9, 96), (47, 68)]
```

## Question 1-1 (1 point)

Write code to answer a workload of range queries using `laplace_mech` and sequential composition. Your solution should have a **total privacy cost of epsilon**.

<!-- grade_id: cell-91ba0020d4ea347e, points: 0, grade: false, solution: true -->
```python
def workload_laplace(workload, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_laplace(random_workload, 1.0)[:4])
```

Output:
```
Cell In[5], line 5
      1 def workload_laplace(workload, epsilon):
      2     # YOUR CODE HERE
      3     raise NotImplementedError()
----> 5 print('First 4 answers:', workload_laplace(random_workload, 1.0)[:4])

Cell In[5], line 3, in workload_laplace(workload, epsilon)
      1 def workload_laplace(workload, epsilon):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-3fdb7419fcfc0cea, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_laplace(random_workload, 1.0))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 50
assert np.mean(errors) < 200
```

## Question 1-2 (1 point)

Write code to answer a workload using `laplace_mech_vec` - the version of the Laplace mechanism for **vector-valued** queries. Your solution should *not* use sequential composition, and should have a total privacy cost of `epsilon`.

*Hint*: remember to use L1 global sensitivity.

<!-- grade_id: cell-e4e3cdb5562933b8, points: 0, grade: false, solution: true -->
```python
def workload_laplace_vec(workload, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_laplace_vec(random_workload, 1.0)[:4])
```

<!-- grade_id: cell-68af547658adcdcd, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_laplace_vec(random_workload, 1.0))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 50
assert np.mean(errors) < 200
```

## Question 1-3 (1 point)

In 2-5 sentences, answer the following:
- Did the two solutions differ in terms of their accuracy?
- How do they differ in terms of their use of composition properties of differential privacy?

<!-- grade_id: cell-0da21812c3f96147, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2-1 (1 point)

Write code to answer a workload using `gaussian_mech_vec` - the version of the Gaussian mechanism for vector-valued queries. Your solution should not use sequential composition, should satisfy $(\epsilon, \delta)$-differential privacy, and should have a total privacy cost of (`epsilon`, `delta`).

*Hint*: remember to use L2 sensitivity.

<!-- grade_id: cell-217b044c0cd1697e, points: 0, grade: false, solution: true -->
```python
def workload_gaussian_vec(workload, epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_gaussian_vec(random_workload, 1.0, 1e-5)[:4])
```

<!-- grade_id: cell-7c4c4d2ed403ea72, points: 10, grade: true, solution: false -->
```python
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_gaussian_vec(random_workload, 1.0, 1e-5))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 10
assert np.mean(errors) < 100
```

## Question 2-2 (1 point)

In 2-5 sentences, answer the following:
- Of your solutions in questions 1-3, which ones rely on *sequential composition*?
- Which solution offers the best accuracy?
- Why does this particular solution yield the best accuracy?

<!-- grade_id: cell-d505748348979709, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3-1 (2 points)

Re-implement your solution to Question 2 using *Rényi differential privacy*. Your solution should satisfy $(\alpha, \bar\epsilon)$-RDP.

*Hint*: See https://programming-dp.com/ch8.html for more about *Rényi DP*.

<!-- grade_id: cell-50da3a173b751f3f, points: 0, grade: false, solution: true -->
```python
def workload_gaussian_vec_RDP(workload, alpha, epsilon_bar):
    # YOUR CODE HERE
    raise NotImplementedError()

print('First 4 answers:', workload_gaussian_vec(random_workload, 1.0, 1e-5)[:4])
```

<!-- grade_id: cell-d8eb625008d2985e, points: 10, grade: true, solution: false -->
```python
# TEST CASE
errors = [abs(r_a - l_a) for (r_a, l_a) in zip(real_answers, workload_gaussian_vec_RDP(random_workload, 5, 0.1))]
print('Average absolute error:', np.mean(errors))
assert np.mean(errors) > 10
assert np.mean(errors) < 100
```

## Question 3-2 (1 point)

Implement a function `convert_RDP_ED` to convert from the $(\alpha, \bar\epsilon)$ of Rényi differential privacy to the $(\epsilon, \delta)$ of approximate differential privacy. Your function should also take the desired value of $\delta$.

<!-- grade_id: cell-361a5a41cab1e483, points: 0, grade: false, solution: true -->
```python
def convert_RDP_ED(alpha, epsilon_bar, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

convert_RDP_ED(5, 0.1, 1e-5)
```

<!-- grade_id: cell-98f137a5295dc1e9, points: 10, grade: true, solution: false -->
```python
# TEST CASE
assert convert_RDP_ED(5, 0.1, 1e-5) == 2.9782313662425572
assert convert_RDP_ED(40, 0.1, 1e-5) == 0.39520321705051864
assert convert_RDP_ED(500, 1.0, 1e-5) == 1.02307199491978
assert convert_RDP_ED(40, 1.0, 1e-5) == 1.2952032170505188
```

## Question 3-3 (2 points)

In 2-5 sentences, answer the following:
- Try various values for `alpha` and `epsilon_bar` in `convert_RDP_ED`. At what values do you observe an $(\epsilon, \delta)$ value around $(1.0, 10^{-5})$?
- Try these values for `alpha` and `epsilon_bar` in `workload_gaussian_vec_RDP`. How does the error compare to using `workload_gaussian_vec`?
- Is it useful to use Rényi differential privacy to answer workloads of range queries? Or is regular $(\epsilon, \delta)$-differential privacy just as good?

<!-- grade_id: cell-17fe437fb638503d, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE